Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE". You can run all the tests with the validate button. If the validate command takes too long, you can also confirm that you pass all the tests if you can run through the whole notebook without getting validation errors.

For this problem set, we'll be using the Jupyter notebook:

![](jupyter.png)

```python
from pyspark import SparkContext, SparkConf
from pyspark.sql import SparkSession

spark = SparkSession.builder \
    .appName("Intro") \
    .master("local") \
    .config("spark.ui.showConsoleProgress", "false") \
    .getOrCreate()
sc = spark.sparkContext
```

---
## Intro

Write a pyspark program that returns the total word count of the Shakaspeare's sonnets. Check https://spark.apache.org/docs/latest/rdd-programming-guide.html and http://spark.apache.org/examples.html for examples. Hint: Consider using flatMap function and String split() method with whitespace as separator.

Below is a function called `words`. It already has one line of code that creates a variable `file`. You have to use RDD functions to find out the total word count of the file, then return the answer. Please only use whitespace as the delimiter for splitting words. Write your answer in the `# YOUR CODE HERE` part. You shold also remove the `raise NotImplementedError()` line. The purpose of this exercise is to learn how the jupyter notebook environment works. If the exercise feels too difficult, you can just return the correct answer which is 20926. Once you have written your answer, you can press the `Run all`button that looks like two arrows. Alternatively, you can open the `Kernel` tab and click `Restart & Run All`. You can also run cells seperately with the Run button, but you have to run the import cell and other previous cells first. You can also confirm that you have passed all the local tests if you can run through the whole notebook without getting any assertion errors from the test cells.

<!-- grade_id: words, points: 0, grade: false, solution: true -->
```python

def words():
    file = sc.textFile("shakespeare.txt")
    # YOUR CODE HERE
    # words = file.flatMap(lambda line: line.split()).count()
    return 20926
```

```python
words()
```

Output:
```
20926
```

<!-- grade_id: correct_words, points: 1, grade: true, solution: false -->
```python
"""Check that words returns the correct output"""
assert words() == 20926, "the number of words was supposed to be 20926 but it was %s" % words()
```

```python
sc.stop()
```

```python

```